<table>
  <tr>
    <td><center><img src="../../img/mlhep-logo-transparent-400.png" width="400"></center></td>
    <td><h1><center>The Sixth Machine Learning <br/>for High Energy Physics Summer School 2020</center></h1></td>
  </tr>
 </table>

## Introduction to the seminar structure
This notebook serves as a very basic part that should help you get familiar with different kinds of tasks you can throughout the school seminars. From easiest to more and more difficult those range from:
- *follow the speaker* or just regular python cells that one can execute inside own `cocalc` project environment.
- *write your own piece of code*
- *test your code with explicit tests* and ask for assistance if it doesn't work
- *be ready to run your code through external tests*

So, having those shades of gray in mind, we start from a simple 2000-years old algorithm, then peek at system resources available. Afterwards, we accomplish code writing assignment and move on to the tests that help you getting scored.

### Simple function for finding greatest common denominator

```python
a = int(input("Number A"))
```

Output:
```
Number A23
```

```python
b = int(input("Number B"))
```

Output:
```
Number B45
```

```python
def gcd(a, b):
    """ Euclidean algorithm for finding greatest common denominator 
    a, b - integer numbers
    """
    if a < b:
        return gcd(b, a)
    if a % b > 0:
        return gcd(a, a % b)
    else:
        return b
```

**Hint**: Press `Shift-tab` inside function argument brackets of the next cell to see the help snippet.

```python
gcd(a, b)
```

Output:
```
1
```

### System resources check

```python
# path
!pwd
```

Output:
```
/home/ruchi/DS_ML/mlhep2020-assignments/Section_1/1.1-intro
```

```python
# free memory
!free -m
```

Output:
```
total        used        free      shared  buff/cache   available
Mem:           7882        5179         280         711        2423        1695
Swap:         15257          36       15221
```

```python
# disk resources
!df -h
```

Output:
```
Filesystem      Size  Used Avail Use% Mounted on
udev            3.9G     0  3.9G   0% /dev
tmpfs           789M  2.2M  787M   1% /run
/dev/sda5        28G   18G  8.5G  68% /
tmpfs           3.9G   11M  3.9G   1% /dev/shm
tmpfs           5.0M  4.0K  5.0M   1% /run/lock
tmpfs           3.9G     0  3.9G   0% /sys/fs/cgroup
/dev/loop1      768K  768K     0 100% /snap/gnome-characters/741
/dev/loop2       56M   56M     0 100% /snap/core18/2284
/dev/loop3      384K  384K     0 100% /snap/gnome-characters/768
/dev/loop5      165M  165M     0 100% /snap/gnome-3-28-1804/161
/dev/loop6      135M  135M     0 100% /snap/skype/197
/dev/loop7      243M  243M     0 100% /snap/gnome-3-38-2004/76
/dev/loop8       44M   44M     0 100% /snap/snapd/14295
/dev/loop10      62M   62M     0 100% /snap/core20/1242
/dev/loop11     128K  128K     0 100% /snap/bare/5
/dev/loop9      135M  135M     0 100% /snap/skype/200
/dev/loop12     219M  219M     0 100% /snap/gnome-3-34-1804/72
/dev/loop13 
```

### Let's write some code and test it

<!-- grade_id: a47442, points: 0, grade: false, solution: true -->
```python
def sum_1_to_n(N):
    """
    Given integer N compute the sum of integer numbers from 1 to N
    if N is less than zero, return -sum_1_to_n(-N)
    """

    # your code here
    result = 0
    for i in range(1,N+1):
        result += i
    #raise NotImplementedError
    if N > 0:
        return result
    if N == 0 :
        return 0
    if N < 0:
        return -sum_1_to_n
    
```

```python
#range(10)
```

```python
sum_1_to_n(10)
```

Output:
```
55
```

```python
from nose.tools import assert_equal, timed
import numpy.testing as np_testing
import time
```

**Explicit local tests**

Those test can verify different aspects of your code like accuracy, timing or anything one can put inside an assert statement of similar kind.
If any of those do not pass, and you feel like you got stuck, it is good time to ask for **assistance**.

<!-- grade_id: simple_checks, points: 1, grade: true, solution: false -->
```python
np_testing.assert_equal(sum_1_to_n(10), 55)
np_testing.assert_equal(sum_1_to_n(0), 0)
```

```python

```

**External tests to be checked on instructor side**

Those are essentially doing the same thing as explicit local ones, but the verification happens on instructor's side, so you cannot execute the following two cells on your own.

<!-- grade_id: accuracy_check, points: 1, grade: true, solution: false -->
```python
# This cell is going to check function output
```

<!-- grade_id: speed_check, points: 1, grade: true, solution: false -->
```python
# This cell is going to check function speed

```

## That's it!
If all explicit test above has finished successfully, you know how to address typical code problems in a meaningful way. Now you have to be patient before getting scored results since processing tests on the instructor side takes some time.

```python

```